<font color ='Navy'><center> __BruFace__ </center></font>
<br> MA1 Mechanical Engineering <br>  __Mechanical Vibrations__ <br>  <br> 

***
<font color='Navy'><h1><center>Session 2 : One DOF system</center></h1></font> 
***
<div class="alert alert-block alert-info">
<center><b>Welcome to the second session </b></center>

```python
# don't worry about this line 
# For local use : %matplotlib notebook; for use on jupyterLab: %matplotlib widget
%matplotlib widget
```

# About the Practical sessions
## Python and Jupyter

The practical sessions are organized using "Jupyter Notebooks" a tool to program and study results in the Python programming language. It is not our goal to make you ace Python programmers but we do believe it is a useful skill for a modern day engineer. If you are interested to get started with Python there are many resources available online, we recommend: https://realpython.com/learning-paths/python3-introduction/

For these practical sessions we will use these partially filled in notebooks, so you don't need to worry too much about the nuts and bolts of Python and we focus on the Mechanical Vibrations. We do hope you also get a bit familiar with Python in the progress.

<div class="alert alert-block alert-info">
<center><b>It is important that you run each cell of the notebook. To do so select a cell (it will be highlighted) and press shift+enter on your keyboard or the play button in the menu above.</b></center></div>

__During the practical session we will make use of three very common library in python. Namely, numpy, scipy and matplotlib__ 
<br>
1. NumPy is a fundamental package for scientific computing in Python. It provides powerful matrix manipulation, basic linear algebra and much more
2. SciPy is a collection of mathematical algorithms and convenience functions built on the NumPy extension of Python.
3. Matplotlib is a Python 2D plotting library. <br>
***
NB : a package is just a collection of code (module) that you can import and use. If you want to learn more about a package, you shall go and visit the website of the package. E.g. for SciPy: https://docs.scipy.org/doc/scipy/

Let's import both numpy and matplotlib.pyplot. The `as` in the code below allows us to use the shorthand `np` and `plt` 
for these packages.



```python
import numpy as np 
import matplotlib.pyplot as plt
```

Then when we want to use a function "foo" in the package numpy we have to type np.foo(arguments)

## Making the exercises
During the exercises you will need to complete some of the cells to make the code work. This will require you to implement some Python code. We kept it fairly simple and we tried to provide all secondary code so you can focus on the assignment at hand.

So let us get started with a first example.

## A first example

To get you started let us start with a first exercise. Let us program a function that converts a pulsation $\omega$ into a frequency $f$. We have set you up with the basic structure below.

For this purpose we already set you up with a new function called `convert_pulsation` which takes in one argument `w` which is the pulsation. We also made the function return a value `f`. Can you complete the code (where it says `# your code here`) so the pulsation `w` is converted into a frequency `f` in _Hz_? You should also remove the statement: `raise NotImplementedError`.


**TIP**: you can access the variable $\pi$ using the command `np.pi`

<!-- grade_id: cell-4f543db7372a06d2, points: 0, grade: false, solution: true -->
```python
# cell metadata=autograded-answer
def convert_pulsation(w:float):
    """
    Implement a function that the frequency from a pulsation omega 
    
    Arguments:
    w -- the pulsation 

    Returns:
    f -- the frequency in Hz
    """   
    
    ### BEGIN SOLUTION
    f = w/2/np.pi
    ### END SOLUTION

    return f
```

Now that this function is created we can use it as follows : `f = convert_pulsation(w=np.pi)` to convert a pulsation of $\pi{}$  _rad/s_ into a frequency.

In the next we can check (=assert) whether or not your solution was correct or not. When the cell below runs without errors then you're code is working correctly. Note that you can't change this cell yourself. 

And, yes you can't see all tests, so we are also checking some other things in the background :-)

<!-- grade_id: cell-2deadf25b08b987f, points: 1, grade: true, solution: false -->
```python
# cell metadata=autograder-tests
assert convert_pulsation(np.pi) == 0.5
assert convert_pulsation(2*np.pi) == 1

### BEGIN HIDDEN TESTS
assert convert_pulsation(4*np.pi) == 2
### END HIDDEN TESTS
```

# Exercises

## Exercise 1 : Fundamental equations of a SDOF system

Consider the following one-degree-of-freedom (1 DOF) system :
<img src="./Images/SDOF.png" width=200 height=200 />


<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b>Assignment 1.1:</b> Write the equation of motion in the time domain (Using LaTeX) in the cell below: 

(If you are not familiar with LaTeX you can use: https://www.hostmath.com/ and copy the LaTeX below)

<!-- grade_id: cell-6c65972325f05f72, points: 0, grade: true, solution: true -->
###BEGIN SOLUTION
$$
m \ddot x(t) + c \dot x(t) + k x(t) = f(t)
$$
###END SOLUTION

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b>Assignment 1.2:</b> Code the expression to calculate $\omega_n$ using the properties of the SDOF system into the function `calcW_n`

<!-- grade_id: cell-c58974031f48ce06, points: 0, grade: false, solution: true -->
```python
def calcW_n(k:float,m:float):
    """
    Implement the expression of W_n

    Arguments:
    k, m -- parameters of the SDOF system

    Returns:
    w_n -- the value of w_n, computed using the formula seen in the course
    """   

    ### BEGIN SOLUTION
    w_n = np.sqrt(k/m)
    ### END SOLUTION

    return w_n
    
```

Validation, if the following cell runs without error you have successfully implemented `calcW_n`:

<!-- grade_id: cell-133ed42043650d13, points: 1, grade: true, solution: false -->
```python
# cell metadata=autograder-tests

assert calcW_n(12,3) == 2.0
assert calcW_n(16,4) == 2.0
assert calcW_n(64,4) == 4.0

### BEGIN HIDDEN TESTS
for i in range(5):
    a=np.random.randint(100)
    b=np.random.randint(100)
    assert calcW_n(a,b)==np.sqrt(a/b)
### END HIDDEN TESTS
```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b>Assignment 1.3:</b> Write the expression to calculate $\xi$ using the properties of the SDOF system into the function `calcXi`

<!-- grade_id: cell-ca9770c87b1dadb8, points: 0, grade: false, solution: true -->
```python
def calcXi(m:float, c:float, k:float):
    """
    Implement the expression of Xi
    
    Arguments:
    m,c,k -- parameter of the SDOF system
    
    Returns:
    xi -- the value of Xi, computed using the formula seen in the course
    """   

    ### BEGIN SOLUTION
    xi= c/(2*np.sqrt(m*k))
    ### END SOLUTION
    
    return xi
```

Validation, if the following cell runs without error you have successfully implemented `calcXi`:

<!-- grade_id: cell-6c396b258e66a9bd, points: 2, grade: true, solution: false -->
```python
assert calcXi(1,20,1) == 10.0
assert calcXi(2,20,2) == 5.0
assert calcXi(3,36,3) == 6.0

### BEGIN HIDDEN TESTS
for i in range(5):
    mt=np.random.randint(1,100)
    ct=np.random.randint(1,100)
    kt=np.random.randint(1,100)
    
    assert calcXi(mt,ct,kt) == ct/(2*np.sqrt(mt*kt))
### END HIDDEN TESTS
```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b>Assignment 1.4:</b> Write the expression to calculate $\omega_d$, starting from $\omega_n$ and $\xi$, into the  function `calcW_d`

<!-- grade_id: cell-6b9b5b0d1588099d, points: 0, grade: false, solution: true -->
```python
def calcW_d(w_n:float,xi:float):
    """
    Implement the expression of W_n
    
    Arguments:
    wn,xi -- properties of the SDOF system
    
    Returns:
    w_n -- the value of w_n, computed using the formula seen in the course
    """   

    ### BEGIN SOLUTION
    w_d = w_n*np.sqrt(1-xi**2)
    ### END SOLUTION
    return w_d
```

Validation, if the following cell runs without error you have successfully implemented `calcW_d`:


<!-- grade_id: cell-cc42252bd442c8db, points: 2, grade: true, solution: false -->
```python
assert round(calcW_d(1,0.1),3) == 0.995 # The round(results,3) rounds your results to 3 decimal pounts
assert round(calcW_d(2,0.05),3) == 1.997

### BEGIN HIDDEN TESTS
for i in range(5):
    w_n = np.random.randint(1,100)
    xi = np.random.randint(1,50)/100
    
    assert round(calcW_d(w_n, xi),3) == round(w_n*np.sqrt(1-xi**2),3)
### END HIDDEN TESTS
```

<div class="alert alert-block alert-info">
<b>End of Exercise :</b> We have reached the end this first assignment, you have now implemented three functions: `calcW_n`, `calcXi` and `calcW_d`. These allow us to calculate the properties of any SDOF system.
</div>

## Exercise 2: Impulse response of a SDOF

Let us put the functions that you wrote in Exercise 1 to good use.  Consider a SDOF system, as illustrated in the Figure at the beginning of Exercise 1, with following properties:

- a mass of $1kg$
- a stiffness of $16N/m$
- a damping coefficient of $0.1 Ns/m$.

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b>Assignment 2.1:</b> Calculate $\omega_n$, $\xi$ and $\omega_d$ for the system described above using the functions made in Exercise 1. 

<!-- grade_id: cell-42a2fbb4ea4ecff7, points: 0, grade: false, solution: true -->
```python

### BEGIN SOLUTION
m = 1
c = 0.1
k = 16
### END SOLUTION

w_n = calcW_n(m=m, k=k)
xi = calcXi(m=m, c=c, k=k)
w_d = calcW_d(w_n, xi)

conclusion = f'The system has a natural frequency of {w_n/2/np.pi:.3}Hz, a damping ratio of {xi*100}% and damped natural frequency of {w_d/2/np.pi:.3}Hz'

print(conclusion)
```

We now have determined the modal parameters for the SDOF system. Let us try to calculate the impulse response of the SDOF system. This impulse response is the behaviour how the system responds to an impulse load. 

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b>Assignment 2.2:</b> Implement the impulse response function for SDOF systems in the function `impulse_response`. This is implemented using the equations you built in Exercise 1.

<!-- grade_id: cell-65fb23fa55a4aaa9, points: 0, grade: false, solution: true -->
```python
def impulse_response(m:float,c:float,k:float,t:np.ndarray):
    """
    Implement the SDOF impulse response 
    
    Arguments:
    m, c and k -- Parameter of the system 
                -- m is in kg
                -- c is in N.s.m^-1
                -- k is in N.m^-1
    t          -- Time vector, containing the timesteps 
    
    Returns:
    h -- the value of function h for the timesteps in the Time vector t, computed using 
        expression of the impulse response
    -----------------------------
    tips :  1.call the function defined above to compute Xi and Wn and 
              save them in a local variable
            2.define a variable wd first
            3.use np.exp and np.sin
    -----------------------------
    """

    xi=calcXi(m=m, c=c, k=k)
    
    w_n=calcW_n(m=m, k=k)
    
    w_d= calcW_d(w_n, xi)
    
    ### BEGIN SOLUTION
    h=1/(m*w_d)*np.exp(-xi*w_n*t)*np.sin(w_d*t)
    ### END SOLUTION
    return h
    
```

Validation, if the following cell runs without error you have successfully implemented `impulse_response`. To test the functions we will used some logical tests. Have a look at the tests below, do you understand their motivation?

<!-- grade_id: cell-3355122c2dad6930, points: 3, grade: true, solution: false -->
```python
m = 1
c = 0.1
k = 16

assert impulse_response(m=m, c=c, k=k, t=0) == 0      # at t=0, no response yet, the system is at rest
assert impulse_response(m=m, c=c, k=k, t=1e8) < 1e-18  # After a long time, the system will no longer oscillate

# What happens in absence of damping (c=0) ?
assert  not impulse_response(m=m, c=0, k=k, t=1e8) < 1e-18 # In absence of damping, after a long time the system will still oscillate!
```

Did you understand all of the tests we used above? 
Maybe it helps to understand them if we plot the first 100 _s_ of the impulse response for the system of interest.

```python

t =np.arange(start=0,stop=100,step=0.1) # Defining a time vector with 0.1s step size
h =impulse_response(m=1,k=16,c=0.1,t=t) # Calculating the impulse reponse

# Plotting the result:

fig, ax = plt.subplots(figsize=(7, 5))

# Plot the data
ax.plot(t, h)

# Set labels and grid
ax.set_xlabel('Time (s)')
ax.set_ylabel('h(t)')
ax.grid(which='both', linestyle=':')

# Display the plot
plt.show()
```

We can clearly see the role of $\omega_d$ and $\xi$ in this impulse response $h(t)$. In the figure below we can nicely see how the impulse response is the combination of an undamped sinus (with a frequency $\omega_d$) and an exponential decay (driven by the damping ratio $\xi$).

```python
m = 1
k = 16
c = 0.1

t =np.arange(start=0,stop=100,step=0.1) # Defining a time vector with 0.1s step size
h =impulse_response(m=m,k=k,c=c,t=t) # Calculating the impulse reponse

xi=calcXi(m=m, c=c, k=k)
w_n=calcW_n(m=m, k=k)
w_d=calcW_d(w_n, xi)

# Plotting the result:

# Create a subplot
fig, ax = plt.subplots(figsize=(7, 5))

# Plot the data
ax.plot(t, np.sin(w_d * t) * np.max(h), color='lightgray', label=r'$\sin(\omega_d\cdot{}t)$')
ax.plot(t, np.exp(-xi * w_n * t) * np.max(h), color='firebrick', label=r'$e^{-\xi\cdot{}\omega_d\cdot{}t}$')
ax.plot(t, h, label=r'$h(t)$')

# Set labels, limits, grid, and legend
ax.set_xlabel('Time (s)')
ax.set_ylabel('h(t)')
ax.set_xlim([0, 50])
ax.grid(which='both', linestyle=':')
ax.legend()

# Display the plot
plt.show()
```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b>Assignment 2.3:</b>: We have made an interactive tool to plot how the impulse reponse function varies when you vary $m$, $c$, $k$. Try it out to better understand the interplay of the SDOF properties and it's dynamic behaviour. No programming required. 

```python
import sys
import iplot.interactive_plot_1
#del sys.modules["iplot.interactive_plot_1"]
#import iplot.interactive_plot_1
```

<div class="alert alert-block alert-success">
<b>Suggestion :</b>
Try to find the parameter that leads to a damping coefficient equal to 100%, observe how $\omega_d$ becomes nan (= 'Not a Number') since the assumption used to derive the expression of $\omega_d$ no longer holds (damping coef <1) 
</div>
<div class="alert alert-block alert-info">
<b>End of Exercise :</b> We have reached the end this first assignment, in which you calculated the impulse response of a SDOF system. You should now understand the interplay `m`, `c` and `k` and the modal parameters and how this translates into a different impulse response function.
</div>
    
## Exercise 3 : Forced harmonic response.

Let us again consider our SDOF system. But this time try to understand how it would respond to a harmonic force:

$f(t) = sin(\omega_f*t)$.

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b>Assignment 3.1:</b> Write a function `harmonic_force` that for a defined driving frequency $f$ (in Hz) calculates the harmonic force $f(t)$ for any timestamp $t$.

<!-- grade_id: cell-25a6392a07771bf2, points: 0, grade: false, solution: true -->
```python
def harmonic_force(f:float, t:float):
    """
    Calculate the harmonic force using a driving frequency f for any timestamp t
    
    Arguments:
    f -- driving frequency  in Hz
    t -- timestamp (in s)
    
    Returns:1
    force -- the harmonic force at timestamp t
    
    """
    ### BEGIN SOLUTION
    w_f = f*np.pi*2
    force = np.sin(w_f*t)
    ### END SOLUTION
    return force

```

Validation, if the following cell runs without error you have successfully implemented `harmonic_force`

<!-- grade_id: cell-e779916fcb75a545, points: 1, grade: true, solution: false -->
```python
assert round(harmonic_force(f=1,t=1),1) == 0 
assert round(harmonic_force(f=2,t=1),1) == 0 
assert round(harmonic_force(f=0.5,t=1),1) == 0 
assert round(harmonic_force(f=0.25,t=1),1) == 1 

```

The next step is to calculate the response of our SDOF system $x(t)$ to this harmonic force. To do so we will use Duhamel's integral.

$$ \begin{align}\begin{aligned} x(t)=\int_{0}^{t} f(\tau)h(t-\tau) d\tau 	 \end{aligned}\end{align} $$

Or since $h(t)=0$ and $f(t)=0$ for all t<0, so we can write it as a convolution:

$$ \begin{align}\begin{aligned} x(t)=\int_{-\infty}^{\infty} f(\tau)h(t-\tau) d\tau 	 \end{aligned}\end{align} $$
$$ \begin{align}\begin{aligned} x(t)= f(t) \ast h(t) 	 \end{aligned}\end{align} $$

This is quite simple to implement, we use the `convolve` function from `scipy.signal` package to return the convolution of the outcome of `harmonic_force` and `impulse_response`.

`Tip`: do not forget to multiply the result of the numerical convlution with the time step `dt`.

```python
def harmonic_response(m:float,c:float,k:float,f:np.ndarray,t:np.ndarray):
    """
    Implement the SDOF impulse response 
    
    Arguments:
    m, c and k -- Parameter of the system 
                -- m is in kg
                -- c is in N.s.m^-1
                -- k is in N.m^-1
    t          -- Time vector, containing the timesteps 
    
    Returns:
    h -- the value of function h for the timesteps in the Time vector t, computed using 
        expression of the impulse response
    -----------------------------
    tips :  1.call the function defined above to compute Xi and Wn and 
              save them in a local variable
            2.define a variable wd first
            3.use np.exp and np.sin
    -----------------------------
    """
    from scipy.signal import convolve
    
    f = harmonic_force(f, t)
    h = impulse_response(m=m, c=c, k=k, t=t)
    dt = t[1] - t[0]
    response = convolve(h, f, mode='full')*dt # Convolution of the force and the impulse response
    # multiplying by dt to get the correct amplitude
    
    return response[:len(f)]
```

Let us have a look at how the harmonic response looks like for our SDOF system for a harmonic excitation at 0.05Hz.

```python
# Define the parameters of the system
m = 1
c = 0.1
k = 16
# Define the time vector
t = np.arange(start=0,stop=100,step=0.1)
# Calculate the natural frequency
w_n=calcW_n(m=m, k=k)
# Calculate the response of the system
response = harmonic_response(m=m, c=c, k=k, f=0.05, t=t)
# Prepare the plot
fig, ax = plt.subplots(figsize=(7, 5))

# Plot the harmonic response
ax.plot(t, response)

# Set labels and grid
ax.set_xlabel('Time (s)', fontsize=12)
ax.set_ylabel('Harmonic response x(t)', fontsize=12)
ax.grid(which='both', linestyle=':')

# Display the plot
plt.show()

```

<div class="alert alert-block alert-info">
<b>Remark :</b> In the above figure the transient and the steady state regime can be clearly observed.
The first 3 cycles correspond to the transient regime.
</div>

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b>Assignment 3.2:</b> We have made an interactive tool to plot how the harmonic reponse function varies when you vary the driving frequency $f$. Try it out to better understand the interplay of the SDOF system and the driving frequency. Which types of transient behaviour can you spot? What is the relation with structures natural frequency $\omega_n$? 

```python
import sys
import iplot.interactive_plot_2
#del sys.modules["iplot.interactive_plot_2"]
#import iplot.interactive_plot_2
```

<div class="alert alert-block alert-success">
<b>Suggestion :</b> Vary the excitation frequency and observe how the transient regime behaves: In the transient regime, we have the presence of two harmonics: the natural frequency + the excitation frequency, which for certain frequency values leads to the beating phenomenon.
</div>

In the previous assignment you could clearly see different types of transient behaviour. After this transient behaviour has passed we can observe the steady state behaviour. 

One way to analyse the steady-state behaviour of a system subjected to different harmonic excitations is to work with the frequency domain expression and plotting it in a **Bode diagram**.

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b>Assignment 3.3:</b> Write a function `frequency_response_function` that computes the frequency response function, for harmonic excitations, for any frequency $f$ (in Hz) taking the system properties $m$, $c$ and $k$ as inputs.



<!-- grade_id: cell-e0b915a75a944d06, points: 0, grade: false, solution: true -->
```python
def frequency_response_function(m:float,c:float,k:float,f:np.ndarray):
    """
    Implement the SDOF expression of the harmonic forced response
    
    Arguments:
    m, c and k -- Parameter of the system 
                -- m is in kg
                -- c is in N.s.m^-1
                -- k is in N.m^-1
    f --  vector of frequencies (Hz)
    
    Returns:
    H -- a vector of the value of the function H, computed using 
        expression of the expression of the harmonic forced response
    -----------------------------
    tips : In python the complex i can be expressed using the letter j
    
    -----------------------------
    """
    ### BEGIN SOLUTION
    w = np.pi*2*f
    wn = calcW_n(m=m,k=k)
    xi = calcXi(m=m,c=c,k=k)
    H=1/(m*(wn**2-w**2+2j*xi*w*wn))
    ### END SOLUTION
    return H

```

Validation, if the following cell runs without error you have successfully implemented `frequency_response_function`:




<!-- grade_id: cell-f082c1cb5b46b04b, points: 3, grade: true, solution: false -->
```python
assert isinstance(frequency_response_function(m=1, c=0.1, k=16, f=10), complex) # A FRF function is a complex number
assert round(abs(frequency_response_function(m=1, c=0.1, k=16, f=0.62)),1) == 1.1 
assert round(abs(frequency_response_function(m=1, c=0.1, k=16, f=0.66)),1) == 0.8 

```

As you can tell from the tests above the outcome of the frequency response function is a complex number. Therefore it carries information about both the **Amplitude** and **phase**. A Bode diagram visualizes exactly these elements of a frequency response function. Let us draw a Bode diagram of our SDOF system for the frequency range of 0 to 2 _Hz_.

We already provided a function `bode_plot` that nicely draws the Bode plot.

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b>Assignment 3.4:</b> Use the provided function `bode_plot` to draw the Frequency response function (in the range from 0 to 2 _Hz_) for a system with :

- a mass of  1𝑘𝑔 
- a stiffness of  16𝑁/𝑚 
- a damping coefficient of  0.1𝑁𝑠/𝑚.




<!-- grade_id: cell-1f23211243db875f, points: 0, grade: false, solution: true -->
```python
def bode_plot(f, H):
    """
    Draw a bode plot for the frequency response function H and corresponding frequency vector f.
    
    Arguments:
    H -- a vector of the value of the function H, computed using 
            expression of the harmonic forced response
    f --  vector of frequencies (Hz)
    """
    
    
    fig,ax= plt.subplots(nrows=2,figsize=(8,6))
    ax[0].plot(f,np.abs(H))
    ax[0].set_yscale('log')
    ax[0].set_ylabel('Amplitude ')
    ax[0].set_xlabel('Frequency (Hz)')
    ax[0].grid(which='both', linestyle=':')

    ax[1].plot(f,np.angle(H,deg=True))
    ax[1].set_xlabel('Frequency (Hz)')
    ax[1].set_ylabel('Phase (deg)')
    ax[1].grid(which='both', linestyle=':')


### BEGIN SOLUTION
f= np.arange(0,2,0.002)
H = frequency_response_function(m=m,c=c,k=k,f=f)
bode_plot(f,H)
### END SOLUTION
```

So what does this mean? To illustrate how this relates to the system response to a harmonic force (of a fixed amplitude but variable frequency) we made a small demonstration. Note that despite we are not changing the input amplitude that the output amplitude does vary! Can you maximize the response amplitude?





```python
import sys
import iplot.interactive_plot_3
#del sys.modules["iplot.interactive_plot_3"]
#import iplot.interactive_plot_3
```

So what about the phase difference between the harmonic force and the system's response? To study this let us just look at the first 20 seconds of the signal.

```python
import sys
import iplot.interactive_plot_4
#del sys.modules["iplot.interactive_plot_4"]
#import iplot.interactive_plot_4
```

For our final assignment we will have a look at how the Frequency response function varies for different systems.

```python
import sys
import iplot.interactive_plot_5
#del sys.modules["iplot.interactive_plot_5"]
#import iplot.interactive_plot_5
```

<div class="alert alert-block alert-success">
<b>Remark :</b> Notice how, when you increase the damping, the amplitude of the resonance decreases. However, for frequencies far from the resonance frequency, the transfer function does not change much even for a high damping.
</div>

## Exercise 4 : Dynamic Response to Base Excitation

Consider the following 1 DOF system subjected to base excitation $ x_{0} $:

<img src="./Images/Base_Exitation.png" width=800 height=300 />

In the following cell, the Santa Cruz earthquake (1990) data is imported from a file on the server. Using the package `pandas` imported as the shorthand `pd`, the `read_csv` command is used to read the data from a .csv file. Once imported, the data is stored in a so-called `Dataframe` (hence why we call it `df`). Using the `df.head()` method we show the first five lines of the Dataframe.

```python
import pandas as pd

df = pd.read_csv('santa_cruz.csv', header=0)
df.head()
```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b>Assignment 4.1:</b> Use the `ax.plot` command to visualize the Santa Cruz earthquake (1990). Complete the cell below by plotting `acc` vs `time` in your `df`. Label the axis correctly with the right units.

<!-- grade_id: cell-70c5cb20bfbac842, points: 0, grade: false, solution: true -->
```python
fig, ax = plt.subplots(figsize=(7, 5))

t = df['time (s)']
a = df['acc (m/s2)']

###BEGIN SOLUTION
ax.plot(t, a, label='Santa Cruz earthquake')
ax.legend()
ax.grid(which='both', linestyle=':')
ax.set_xlabel('Time ($s$)')
_ = ax.set_ylabel('Acceleration ($m/s^2$)')
###END SOLUTION
```

In order to check that you plotted correctly the Santa cruz earthquake, store the peak-to-peak value of the signal in the variable `ptp`. Tip : You can compute the peak-to-peak of the signal using `np.ptp`.

<!-- grade_id: cell-a75f820f279bba50, points: 0, grade: false, solution: true -->
```python
ptp = None
###BEGIN SOLUTION
ptp = np.ptp(a)
# Alternative solution: ptp= np.max(a)-np.min(a)
###END SOLUTION
```

<!-- grade_id: cell-efe53711809b8280, points: 1, grade: true, solution: false -->
```python
# cell metadata=autograder-tests
assert ptp, "You didn't set a value for `ptp`"
assert ptp > 0, "The ptp value can not be negative"
### BEGIN HIDDEN TESTS
assert np.isclose(ptp, 7.6518, atol=0.1), "The ptp value you typed is wrong"
### END HIDDEN TESTS
```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b>Assignment 4.2:</b> Complete the function `response_to_base_excitation` that takes into account the 1 DOF system shown in Exercise 4, subjected to base excitation. Use `convolve` command from `scipy.signal` package for the convolution of `impulse_response` and `base_excitation`. 

`Tip`: do not forget to multiply the result of the numerical convolution with the time step `dt`.

<!-- grade_id: cell-73b04d72ab4dfe8a, points: 0, grade: false, solution: true -->
```python
from scipy.signal import convolve

def response_to_base_excitation(m:float,c:float,k:float,x_o_acc:np.ndarray,t:np.ndarray):
    """
    Implement the SDOF base excitation response 
    
    Arguments:
    m, c and k -- Parameters of the system 
                -- m is in kg
                -- c is in N.s.m^-1
                -- k is in N.m^-1
    t           -- Time vector
    x_o_acc         -- Base acceleration 
    
    Returns:
    response -- the response of the 1DOF to base excitation
    -----------------------------
    """

    ###BEGIN SOLUTION
    f = -m * x_o_acc
    h = impulse_response(m=m, c=c, k=k, t=t)
    dt = t[1] - t[0]
    response = convolve(h, f, mode='full')*dt  # Multiply with dt to get the right amplitude
    ###END SOLUTION
    
    return response[:len(f)]
```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b>Assignment 4.3:</b>  Compute the response of 1 DOF system to Santa Cruz earthquake, save your results in a variable called `response_xo`. Use the following properties of the system: 
- a mass of $1 kg$
- a stiffness of $10 N/m$
- a damping coefficient of $0.1 Ns/m$.

<!-- grade_id: cell-243b8b8180623d8e, points: 0, grade: false, solution: true -->
```python
t = None # time vector of Santa Cruz earthquake
x_o_acc = None # acceleration of Santa Cruz earthquake
response_xo = None # response to Santa Cruz earthquake

###BEGIN SOLUTION
# Define the parameters of the system
m = 1
c = 0.1
k = 10
# Define the time vector and base acceleration
t = df['time (s)']
x_o_acc= df['acc (m/s2)']
# Calculate the response of the system
response_xo = response_to_base_excitation(m=m, c=c, k=k, x_o_acc=x_o_acc, t=t)
###END SOLUTION

fig, ax = plt.subplots(nrows=1,figsize=(7, 5))
ax.plot(t, response_xo)
ax.set_xlabel('Time (s)', fontsize=12)
ax.set_ylabel('Response (m)', fontsize=12)
ax.grid(which='both', linestyle=':')
```

Let us do some assertions to check if you implemented everything correctly.

<!-- grade_id: cell-eb1127458791780c, points: 5, grade: true, solution: false -->
```python
assert len(x_o_acc) == len(t), "Input acceleration and time vector must have the same length"
assert len(response_xo) == len(x_o_acc), "Input acceleration and response must have the same length"

### BEGIN HIDDEN TESTS
assert np.isclose(np.max(response_xo), 0.04, atol=1e-2), "The maximum value of the response is not correct"
assert np.isclose(np.mean(np.square(response_xo)),6.33e-5, atol=1e-4), "The ms in your response are incorrect"
assert np.isclose(np.std(response_xo),0.008,rtol=1e-1) 
### END HIDDEN TESTS
```

Now, let us say we want to find the worst possible spring stiffness for our 1DOF system that yields the highest response to to the Santa Cruz earthquake. 


<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b>Assignment 4.4:</b> Find the worst stiffness value `k_worst` of the 1 DOF system by plotting the peak-to-peak values for   springs with a range of different stiffnesses from 1 to 100N/m, stored as `k_range`. The worst spring would result in higher oscillations with the largest peak-to-peak value. 
    
    TIP: You can use `np.linspace` to define a range of values

<!-- grade_id: cell-40a99bb498040a97, points: 0, grade: false, solution: true -->
```python
ptp_values = []  
k_range = []
###BEGIN SOLUTION
k_range = np.linspace(1, 100, 200)  # Range of spring stiffness values
for k in k_range:
    response = response_to_base_excitation(m, c, k, x_o_acc, t)  # Compute response
    ptp_values.append(np.ptp(response))  # Compute peak-to-peak
max_index = np.argmax(ptp_values)
k_worst= k_range[max_index]
###END SOLUTION

assert len(ptp_values) == len(k_range), "You didn't have an equal amount of results as the number of springs in `k_range`"


#%% Let us plot the results (ptp vs. Spring stiffness)
fig, ax = plt.subplots(nrows=1,figsize=(7, 5))
ax.plot(k_range, ptp_values, label="Peak-to-Peak", color="tab:blue")

ax.set_xlabel("Spring Stiffness (N/m)", fontsize=12)
ax.set_ylabel("Peak-to-Peak (m)", fontsize=12)

ax.grid(which='both', linestyle=':')
ax.set_xlim(left=0, right=max(k_range))
_ = ax.legend()
```

Write in the cell below the value of the worst spring stiffness you found in your plot.

<!-- grade_id: cell-8d9c8bd14ebe13e2, points: 0, grade: false, solution: true -->
```python
k_worst = None
###BEGIN SOLUTION
k_worst = 23.38
###END SOLUTION
```

Validation, if the following cell runs without error you have successfully identified `k_worst`:

<!-- grade_id: cell-6ca76e303fe1521c, points: 3, grade: true, solution: false -->
```python
# The next two assertions just chek if you followed the assignment
assert k_worst, "You didn't define a value for `k_worst`"
assert isinstance(k_worst, float), f"`k_worst` should be a value (float), in your case it is a `{type(k_worst)}`"
assert k_worst > 1, "`k_worst` should be greater than 1N/m"
assert k_worst < 100, "`k_worst` was limited to a maximum of 100N/m"
### BEGIN HIDDEN TESTS
assert np.isclose(k_worst, 23.11, atol=2), "This is not the worst K"
### END HIDDEN TESTS
```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

<b>Assignment 4.5:</b> Plot the worst-case response of the 1 DOF system subjected to the Santa Cruz Earthquake. Store the system response with `k_worst` as `response_xo_k_worst`. Compare the results to the original response. Use the `ax.plot` command and different colors for each repsonse for better visualization.

<!-- grade_id: cell-cc24f55127804eaa, points: 0, grade: false, solution: true -->
```python
response_xo_k_worst = None # Compute the response_to_base_excitation with k_worst

###BEGIN SOLUTION
response_xo_k_worst = response_to_base_excitation(m, c, k_worst, x_o_acc, t) 

fig, ax = plt.subplots(figsize=(7, 5))
ax.plot(t, response_xo, color='red', label='k = 10 N/m', zorder=2)
ax.plot(t, response_xo_k_worst, label=f"k = {k_worst:.2g} N/m",zorder=1)
ax.set_xlabel('Time (s)')
ax.set_ylabel('Response (m)')
ax.grid(which='both', linestyle=':')
_ = ax.legend()
###END SOLUTION
```

Validation, if the following cell runs without error you have successfully computed `response_xo_k_worst`:

<!-- grade_id: cell-7add9c288983f4fb, points: 5, grade: true, solution: false -->
```python
assert response_xo_k_worst is not None, "You didn't define `response_xo_k_worst`"
assert len(x_o_acc) == len(t), "Input acceleration and time vector must have the same length"
assert len(response_xo_k_worst) == len(x_o_acc), "Input acceleration and response must have the same length"
assert np.ptp(response_xo_k_worst) > np.ptp(response_xo), "Your worst case result was not worse than the original system, check your results"

### BEGIN HIDDEN TESTS
assert np.isclose(np.max(response_xo_k_worst), 0.14, atol=1e-2), "The maximum value of the response is not correct"
assert np.isclose(np.mean(np.square(response_xo_k_worst)),0.00106, atol=1e-4), "The mean square in your response is incorrect"
assert np.isclose(np.std(response_xo_k_worst),0.033,rtol=1e-1), "The SDV in your response is incorrect"
### END HIDDEN TESTS
```

<div class="alert alert-block alert-info">
    This concludes session 2, don't forget to <b>VALIDATE</b> and <b>SUBMIT</b> your assignment.
</div>