

<font color ='Navy'><center> __BruFace__ </center></font>
<br> MA1 Mechanical Engineering <br>  __Mechanical Vibrations__ <br>
***
<font size="8" color='Navy'><center>Session 1 : Introduction</center></font> 
***
<div class="alert alert-block alert-info">
<center><b>Welcome to this first session, the main objective is to familiarize with jupyter notebooks and the way these sessions are organized. At the same time we will give you an easy start with python. </b></center>
 </div>  

# <u>Table of contents</u>
1. [Variables in Python](#var)
2. [Functions in Python](#func)
3. [List manipulation](#list)
4. [Importing libraries](#lib)
5. [Logic](#logic)
6. [Loops](#loop)
7. [Visualization (plotting)](#plot)


# Welcome

The practical sessions are organized using "Jupyter Notebooks" a tool to program and study results in the Python programming language. It is not our goal to make you ace Python programmers but we do believe it is a useful skill for a modern day engineer. If you are interested to get started with Python there are many resources available online, we recommend: https://realpython.com/learning-paths/python3-introduction/

For these practical sessions we will use these partially filled in notebooks, so you don't need to worry too much about the nuts and bolts of Python and we focus on the Mechanical Vibrations. We do hope you also get a bit familiar with Python in the progress.


# Introduction to Python

This first session is a low-level entry to Python, so if you already have some experience with python these exercises are going to be easy and you can finish early

## Jupyter 
Jupyter notebook allows a combination of code and text in a single document.
You are seeing the Jupyter notebook right now.
Some of the characteristics of Jupyter notebook:

* cell based programming (similar to Mathematica),
* running each individual cell,
* includes **Markdown** and **Code** type cells. 

Jupyter notebooks can be run and edited within, e.g., Visual Studio Code or in a web browser.

<div class="alert alert-block alert-info">
<center><b>It is important that you run each cell of the notebook. To do so select a cell (it will be highlighted) and press shift+enter on your keyboard or the play button in the menu above.</b></center></div>

**Cell running order is important, the cell run order is indicated at the top left of the cell In [q]**

## Making the exercises
During the exercises you will need to complete some of the cells to make the code work. This will require you to implement some Python code. We kept it fairly simple and we tried to provide all secondary code so you can focus on the assignment at hand.

So let us get started by introducing variables.
## Basics of Python <a class="anchor" id="var"></a>

### Variables

Like most programming languages Python allows you to store values (whether they are a single value, a vector or a matrix) in variables. In the cell below we will store a value of 6 into the variable `my_first_variable`

**Reminder**: To run a cell, press **Shift+Enter** or the "(Play) Execute" button above. You don't need to press 'Validate'.



```python
# this is a 'Code' cell
my_first_variable = 6 #
```

Now that this value is stored I can access it in any other cell.  The next cell will print the value of `my_first_variable`

```python
print(my_first_variable)
```

We can do the expected arithmic operations with our variable. E.g. in the next cell we'll multiply `my_first_variable` by 2 and store the result into `two_times_my_first_variable`

```python
two_times_my_first_variable = 2*my_first_variable
print(two_times_my_first_variable)
```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

**Assignment 1.1** : Define a variable named `the_answer_to_everything` and set its value to 42. Define another variable `my_third_variable` which is the product of `my_first_variable` and `the answer to everything`.

**IMPORTANT** Type your code where it states `your code here`. Don't forget to remove `raise NotImplementedError`

<!-- grade_id: cell-d2ee8d2e2849f0fb, points: 0, grade: false, solution: true -->
```python
# cell metadata=autograded-answer

# YOUR CODE HERE
raise NotImplementedError()

```

The next cell will check if you did everything correctly. If you get an error when running the next cell, then you probably made a mistake in the previous cell.

<div class="alert alert-block alert-info">
<center><b>These validation cells are an important part of all of the exercises. You need to succesfully pass them so the code continues to work well.</b></center></div>

<!-- grade_id: cell-829fdce98e273867, points: 1, grade: true, solution: false -->
```python
# cell metadata=autograder-tests
assert the_answer_to_everything == 42

```

### Functions <a class="anchor" id="func"></a>

A key construction in python is the use of functions. They allow you to repeat a calculation without having to type the code every single time. Once you define a function you can keep on using it, no need to repeat your code all of the time.

The figure below illustrates the basic layout of a function. 

<center>

<img width=800 src="./images/function_py.png">
<br>
<font size="3" color="3e5d81ff" face="garamond"><b> How to define a function</b></font>

</center>


In the example below we'll make a function `power_m` that calculates the power `m` of a variable.

**NOTE** It is important to run a cell that defines a function, like the one below, even when it doesn't return a result.

```python
def power_m(x, m):
    """
    This function calculates x^m 
    
    returns : a number
    """
    result =  x**m
    return result
```

Now that we have defined our function in the cell above, we can now use it throughout the notebook.  E.g. we can use it to calculate how much is the `my_first_variable` to the power 2 (it should be 6^2 = 36).

```python
print(power_m(x=my_first_variable, m=2))
```

<div class="alert alert-block alert-info">
<center><b>Something unique and essential to Python is the correct use of identation (using tabs). Python is very strict in this. So when defining a function you need to have always a tab in front of your code. When you make a mistake you'll get an `IndentationError`.</b></center></div>

For instance this example is wrong and will throw an `IndentationError: expected an indented block`:
``` python
def my_bad_example(a,b):
result = a+b
return result
```

You should have:
``` python
def my_bad_example(a,b):
    result = a+b
    return result
```

Now we learned that it is time for you to try it!

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

  
**Assignment 1.2** Complete the function `a_complex_calculation` which takes 3 values (x,y,z) and computes $(x-y)/z$ 

<!-- grade_id: cell-dbcca58d47619e60, points: 0, grade: false, solution: true -->
```python
# cell metadata=autograded-answer

def a_complex_calculation(x, y, z):
    """
    This function calculates (x-y)/z
    
    returns : a number
    """
    # YOUR CODE HERE
    raise NotImplementedError()

    return result
```

<!-- grade_id: cell-8c473b5842449c7b, points: 1, grade: true, solution: false -->
```python
# cell metadata=autograder-tests

assert a_complex_calculation(10, 5, 2) == 2.5
assert a_complex_calculation(10, 10, 2) == 0
```

Let us check you did that correctly, notice how we use the same function multiple times, we are recycling our code. 

### List Basics <a class="anchor" id="list"></a>

Lists are one of the most versatile and commonly used data structures in Python. They allow you to store multiple items in a single variable. Here are some key points about lists:

1. **Creating a List**: Creating a list is as simple as putting different comma-separated values between square brackets.

    ```python
    my_list = [1, 2, 3]
    mixed_list = [1, "Hello", 3.4]
    ```

2. **Accessing Elements**: Accessing elements of a list is done by using the index of the element in the list. The index of the first element is 0, the second element is 1, and so on.

    ```python
    my_list = [1, 'Mechanical', 3]
    print(my_list[0])  # 1
    print(my_list[1])  # Mechanical
    ```

3. **Negative Indexing**: Negative indexing means beginning from the end, -1 refers to the last item, -2 refers to the second last item, and so on.

    ```python
    my_list = [1, 'Mechanical', 3]
    print(my_list[-1])  # 3
    ```

4. **Slicing a List**: Slicing a list is done by using the colon operator.

    ```python
    my_list = [5, 5, 6, 6, 7, 8, 9, 10.5]
    print(my_list[2:5])  # [6, 6, 7]
    print(my_list[:5])   # [5, 5, 6, 6, 7]
    print(my_list[5:])   # [8, 9, 10.5]
    print(my_list[-5:-2])  # [6, 7, 8]
    print(my_list[::2])  # [5, 6, 7, 9]
    print(my_list[::-1])  # [10.5, 9, 8, 7, 6, 6, 5, 5]
    ```

5. **Modifying Elements**: You can modify elements of a list by assigning a new value to a specific index.

    ```python
    my_list = [1, 'Mechanical', 3]
    my_list[1] = 'Vibrations'
    print(my_list)  # [1, 'Vibrations', 3]
    ```

6. **List Methods**: Python lists come with a variety of useful methods.

    ```python
    my_list = [1, 2, 3]
    my_list.append(4)          # [1, 2, 3, 4]
    my_list.insert(1, 5)       # [1, 5, 2, 3, 4]
    my_list.remove(3)          # [1, 5, 2, 4]
    my_list.pop()              # [1, 5, 2]
    print(len([1, 2, 3]))      # 3
    print([1, 2, 3] + [4, 5, 6])  # [1, 2, 3, 4, 5, 6]
    ```



<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

  
**Assignment 1.3** 
In this assignment, you will work with Python lists. You will practice creating, accessing, and manipulating lists, as well as performing operations on the data they contain.
Tasks:
- First make a list named `mixed_list` with the following elements: 1. , 2. , 'hello', 3. ,'world' ,4., 'hi' ,5.  , '$'
- Second pop the last element of `mixed_list` and store it in a variable named `last_element`  
- Third append the value of 6. to `mixed_list` 
- Fourth insert the value of 0.5 at the second position of `mixed_list`
 


<!-- grade_id: cell-409b8c6f0b4618ec, points: 0, grade: false, solution: true -->
```python
# cell metadata=autograded-answer
mixed_list = []
last_element = None 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0388d9fd501c5760, points: 1, grade: true, solution: false -->
```python
# cell metadata=autograder-tests

assert len(mixed_list) == 10

```

<div class="alert alert-block alert-info">
<p>While Python's built-in tools are powerful, they can be limited for complex scientific and numerical computations. Python's true strength lies in its vast ecosystem of open-source libraries. By importing external packages, we can significantly expand Python's capabilities.</p>
<p>One of the most important packages for scientific computing in Python is <strong>NumPy</strong> (Numerical Python). NumPy provides:</p>
<ul>
  <li>A powerful N-dimensional array object</li>
  <li>Sophisticated (broadcasting) functions</li>
  <li>Useful linear algebra, Fourier transform, and random number capabilities</li>
</ul>
<p>NumPy's core is the <code>ndarray</code>, an N-dimensional array that allows for efficient computation on large datasets. It forms the foundation for many other scientific Python libraries.</p>
<p>To use NumPy, you first need to import it:</p>

### Import statement in python  <a class="anchor" id="lib"></a>

One of the essential parts of python is the use of the `import` statement. This command allows us to import functions from so-called packages so we doen't have to calculate them ourselves! 

**NOTE:** This is a bit different than e.g. MATLAB where you don't have to explicitely import a function. 

To demonstrate have a look at the following cell. We start by importing a function `sqrt` from the popular package `numpy`. The `sqrt` command allows us to compute the squareroot of a number. E.g. in the cell below it is used to calculate the square root of 9. 




```python
from numpy import sqrt 
print(sqrt(9))
```

We can also import `numpy` as a whole, importing all its functions. Often the package name is abbreviated. In the cell below we import `numpy` and abbreviate it as `np`.  Note that we now can access the same `sqrt` function using `np.sqrt`.

```python
import numpy as np 
print(np.sqrt(9))
```

Numpy is home to many functions, some of which we will regularly use in these exercise sessions. E.g. you can import the constant $\pi$ from numpy (`pi`). 

**Assignment 1.4:** Complete following function `area_circle` that calculates the surface area of a circle, to do so import `pi` from numpy.

<!-- grade_id: cell-599aa95c80fc5049, points: 0, grade: false, solution: true -->
```python
# cell metadata=autograded-answer

def area_circle(radius):
    """
    Calculate the surface area of a circle with radius `radius`
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return area
```

Let us test you did that correctly, let us check your code for the calculation for two simple circles with radii 10 and 4.

<!-- grade_id: cell-f049e01bb827faa7, points: 0, grade: true, solution: false -->
```python
# cell metadata=autograder-tests

assert round(area_circle(10),5) == 314.15927
assert round(area_circle(4),5) == 50.26548

```

### Arrays <a class="anchor" id="mat"></a>
In these exercise sessions you will be using arrays at regular occasions. They serve to represent vectors and matrices. Arrays are provided throug `numpy.array`.  As we already imported numpy as `np` before we can access them directly from `np.array`.


The easiest way to define a matrix is to start from a comma-seperated lists ([]) in between square brackets ([]) and convert that to a np.array.  Each row is a new item in the list. So to define a matrix with `n` rows

``` python
np.array([[row_1],[row_2],...,[row_n]])
```
E.g. let us define a 2 by 3 matrix:

```python
my_first_matrix = np.array([[1,2,3],[4,5,6]])
print(my_first_matrix)
```

Or similarly some vectors

```python
my_first_vector = np.array([[5],[6],[7]])
print(my_first_vector)
```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">


**Assignment 1.5**: Define a column vector with three elements [8, 9, 10] and call it `my_second_vector`.

<!-- grade_id: cell-6ae697a58ea2c44e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Let us check if you did everything correctly by doing some simple checks. 

**TIP**: try to understand what the tests are doing.

<!-- grade_id: cell-87c8d05fe380cb04, points: 1, grade: true, solution: false -->
```python
assert np.shape(my_second_vector) == (3,1) # check the shape of my_second_vector  it was a row-vector with three elements
assert np.sum(np.power(my_second_vector, 2)) == 245 # check the norm of the vector

```

Now that we've defined our vectors we can do all sorts of operations, such as transposing `np.transpose(a)`. Including the vectorial product (`a*b`) and the dot product (`np.dot(a,b)`) of two vectors `a` and `b`.

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">
    
**Assignment 1.6** : complete the function `vector_product` below that calculates ${\bf a}*{\bf b}^T$ for two vectors ${\bf a}$ and ${\bf b}$. 


<!-- grade_id: cell-3799e61b36006231, points: 0, grade: false, solution: true -->
```python
def vector_product(a,b):
    """
    Calculates a*b^T
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return result
```

Let us test your outcome, again try to understand the nature of the tests. To do so we'll use the existing vectors `my_first_vector` and `my_second_vector` to calculate a new matrix `my_calculated_matrix`. Note how we use `np.shape(my_calculated_matrix)` to check the dimensions of the matrix. 

<!-- grade_id: cell-f4c334477e86295f, points: 1, grade: true, solution: false -->
```python
my_calculated_matrix = vector_product(my_first_vector, my_second_vector)
assert np.shape(my_calculated_matrix) == (3,3) # Is the shape of my result correct?
assert my_calculated_matrix[0,0] == 40 # Is the first element on diagonal equal to 40?
assert my_calculated_matrix[0,1] == 45 # Is the element on (0,1) equal to 45?

```

Did you spot in the tests above how we accessed particular elements in the matrix `my_calculated_matrix`? Indeed using the square brackets we can reach any element (i,j) in the matrix. 

``` python
my_calculated_matrix[i,j]
```
<div class="alert alert-block alert-info">
<center><b>Something counterintuitive in Python for most people is that the indices start with 0 and not 1! </b></center></div>

So if you want to have element on the first row and first column you'll need to use `my_matrix[0,0]`. **Remember this!!** It is a very common mistake for people new to Python. This also means that the last element of a n-by-n matrix is found on `my_matrix[n-1,n-1]`

You can also access an entire row or column using `:`. E.g. to reach the second column of `my_calculated_matrix` you can do `my_calculated_matrix[:,2]`.

<div class="alert alert-block alert-info">
<center><b>Vectors are treated a bit differently in numpy than in MATLAB, which handles matrices more intuitively, and even to experienced people it is easy to make mistakes. So make sure you regularly check your intermediate steps (e.g. using print). </b></center></div>

### Generating Arrays with 'numpy as np' using `np.linspace` and `np.arange`
`np.linspace(start, stop, num)` returns `num` evenly spaced samples, calculated over the interval `[start, stop]`. E.g. `np.linspace(0, 10, 5)` returns `array([ 0.,  2.5,  5.,  7.5, 10.])`.

`np.arange(start, stop, step)` returns evenly spaced values within a given interval. E.g. `np.arange(0, 10, 2)` returns `array([0, 2, 4, 6, 8])`.

Note that `np.linspace` includes the stop value, while `np.arange` does not.

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

**Assignment 1.7** : Complete the function `generate_time_vector` that generates a time vector from `start_time` (included) to `end_time` (not included)  with a sampling frequency of `f`.

<!-- grade_id: cell-dda6a67d6b8a3cb8, points: 0, grade: false, solution: true -->
```python
import numpy as np  
def generate_time_vector(start_time, end_time, f):
    """
    Generates a time vector from start_time to end_time with a step of time_step
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return time_vector
```

<!-- grade_id: cell-424082ded59455bd, points: 1, grade: true, solution: false -->
```python
assert generate_time_vector(0, 1, 100).shape == (100,)
assert generate_time_vector(0, 1, 100)[0] == 0
assert generate_time_vector(0, 1, 100)[99] == 0.99
assert np.sum(generate_time_vector(0, 10, 1)) == 45
```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

**Assignment 1.8** : Make an np.array `my_array` from 0 to 360 with a spacing of 15.
Then in another array `my_second_array` inverse the order of the elements of `my_array`.
Then in another array select the 10 last elements of `my_second_array` and store them in `my_third_array`.
**TIP** : use the `np.arange` function to generate the array and for the rest google is your friend.

<!-- grade_id: cell-843bccfd02a1d133, points: 0, grade: false, solution: true -->
```python
my_array = None
my_second_array = None
my_third_array = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c1395daba8cd00a5, points: 1, grade: true, solution: false -->
```python
assert my_array[0] == 0
assert len(my_array) == 25
assert my_array[1] == 15
assert my_third_array[0] == 135
assert my_third_array[-1] == 0
assert len(my_third_array) == 10
assert my_second_array[0] == 360
```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

**Assignment 1.9** : Build a function that make a vector rotation of an angle $\theta$ of a vector $\bf{a} \in \mathbb{R}^3$. The rotation is done around the z-axis, so the vector $\bf{a}$ is rotated in the xy-plane. The function should return the rotated vector $\bf{b}$.\
The function should be called `rotate_vector_xy` and get two arguments, the vector $\bf{a}$ and the angle $\theta$ in degrees.\
***TIP*** : Matrix multiplication in numpy is done using `np.dot(A,B)` or `A@B` where A and B are matrices.\
While elementwise multiplication is done using `np.multiply(A,B)` or `A*B` where A and B are arrays or matrices.\
\
The rotation matrix is given by:
$
\begin{bmatrix}
\cos(\theta) & -\sin(\theta) & 0 \\
\sin(\theta) & \cos(\theta) & 0 \\
0 & 0 & 1
\end{bmatrix}
$

<!-- grade_id: cell-29809cac010541e3, points: 0, grade: false, solution: true -->
```python
def rotate_vector_xy(vector, angle):
    """
    Rotates a 2D vector by angle degrees
    : vector : np.array
    : angle : float in degrees
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
    return rotated_vector 
```

<!-- grade_id: cell-8eb1396f74acf8f5, points: 1, grade: true, solution: false -->
```python
assert np.isclose(rotate_vector_xy(np.array([1,0,0]), 90), np.array([0,1,0])).all()
assert np.isclose(rotate_vector_xy(np.array([1,0,0]), 180), np.array([-1,0,0])).all()
assert np.isclose(rotate_vector_xy(np.array([1,0,0]), 270), np.array([0,-1,0])).all()
assert np.isclose(rotate_vector_xy(np.array([1,0,0]), 360), np.array([1,0,0])).all()


```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

**Assignment 1.10** : In this assignment, you will work with a predefined 5x5 magic square matrix named  `magic_matrix `. Your task is to extract specific rows, columns, and diagonal elements from this matrix. The manipulation you will learn are refered to Slicing and striding
- Extract the first row of the matrix and store it in a variable named  `first_row `.
- Extract the last row of the matrix and store it in a variable named  `last_row `.
- Extract the main diagonal (from top-left to bottom-right) and store it in a variable named  `diagonal `.
- Extract the first and the last row, and store in a variable named  `first_and_last_row`.
- Extract the first two elements of the last column and store them in a variable named  `first_two_elts_of_last_column `.
- Extract the elements of magic_matrix at every other row and every other column (Striding) (i.e., elements at positions [0, 0], [0, 2], [2, 0], [2, 2], etc.) and store them in a variable named `strided_matrix`. Tips : for 

***Note***: A magic square is a special type of square matrix where the sum of the numbers in each row, each column, and both main diagonals is the same. For a 5x5 magic square, this sum is given by the formula $Cte = \frac{N \times (N^2 +1)}{2}$, where N is the dimension of the square.

**TIP:** 

- **Slicing Notation:** Use the slicing notation `start:stop:step` in NumPy to achieve striding. For example, `matrix[::2]` would select every second element along a dimension. The `::2` indicates that the step size is 2, meaning it skips every other element.
- 
  
- **Default Values:** If `start` or `stop` are not specified, Python assumes you want to start at the beginning or go until the end of the array. For example, `matrix[1::2]` will start from the second element (index 1) and take every second element until the end.

- **Striding in Both Dimensions:** You can apply striding across multiple dimensions by separating them with a comma. For instance, `matrix[::2, ::2]` will stride across both rows and columns.

**TIP:** Refer to the [NumPy Indexing Documentation](https://numpy.org/devdocs/user/basics.indexing.html) for more information on slicing and striding.


<!-- grade_id: cell-142bb598afc598fa, points: 0, grade: false, solution: true -->
```python
magic_matrix = np.array([
    [17, 24, 1, 8, 15],
    [23, 5, 7, 14, 16],
    [4, 6, 13, 20, 22],
    [10, 12, 19, 21, 3],
    [11, 18, 25, 2, 9]
])
first_row=None
last_row=None
diagonal=None
first_and_last_row=None
first_two_elts_of_last_column= None
strided_matrix = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6bb7717d4bb52d15, points: 1, grade: true, solution: false -->
```python
assert np.sum(first_row) == 65
assert np.sum(last_row) == 65
assert np.sum(diagonal) == 65
assert np.sum(first_and_last_row) == 130
assert np.sum(first_two_elts_of_last_column) == 31
assert np.sum(strided_matrix) == 117

```

## Logic <a class="anchor" id="logic"></a>

### if - elif - else

No programming language can survive without basic logical functions. While there are several possible solutions in Python you can do a lot of programming with the `if`,`elif` (= else if),`else` commands.


```python
if your_value == 0 : # Check if your_value equals 0 
    print('your value is equal to zero')
elif your_value > 1 :# If not equal to 0 then check if your_value is perhaps greater than 1
    print('your value is greater than 1')
else : # In all other scenarios do
    print('your value is not equal to zero and not greater than 1')
```
You can have as many `elif` statements as you need.

**NOTE** Similar to functions everything inside the `if`, `elif`and `else` statements needs to be indented! 

### loops <a class="anchor" id="loop"></a>

For loops are very important concept in programming and they make our live easy. They are a bit different in Python than in MATLAB but they are also bit more intuitive. 

For instance to loop over all values from 0 to 9, we first use `np.arange` to make a vector filled with 0 to 9 and then use `for` to loop over it:

``` python
for i in np.arange(10):
    print(i)
```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

**Assignment 2.1** : Finish the function `find_first_value_greater_than(matrix, a)` which returns the first value (working row by row, [0,0]->[0,1]->...->[0,n]->[1,0]->...->[n-1,n-1] ) in a matrix that is greater than the value `a`.

**BIG Tip** : you can use the `return` command in any point of your function, when `return` is reached the function terminates and returns the corresponding value

<!-- grade_id: cell-250a072c3cf366ea, points: 0, grade: false, solution: true -->
```python
def find_first_value_greater_than(matrix, a):
    # YOUR CODE HERE
    raise NotImplementedError()

```

Let us test your function using `my_calculated_matrix`

<!-- grade_id: cell-d185a2d17e15f90b, points: 1, grade: true, solution: false -->
```python
assert find_first_value_greater_than(my_calculated_matrix, 41) == 45
assert find_first_value_greater_than(my_calculated_matrix, 51) == 54
assert find_first_value_greater_than(my_calculated_matrix, 60) == 63 # Not 60

```

# Plotting

To end this session we'll give a quick introduction to plotting results. In python there are several packages for plotting, but the most commonly used is `matplotlib.pyplot`, typically abbreviated as `plt`.  
There are multiple ways in python to make a figure, we will use the `plt.subplots` function to make a figure and axes. 

``` python
import matplotlib.pyplot as plt
x = np.linspace(0,2*np.pi,100)
y = np.sin(x)
fig, ax = plt.subplots()
ax.plot(x,y)
ax.set_xlabel('x')
ax.set_ylabel('y')
ax.set_title('A simple plot')
```
here we use the `np.linspace` function to generate a vector `x` with 100 elements from 0 to $2\pi$ and then calculate the sine of `x` to get `y`. We then plot `y` against `x` using `ax.plot(x,y)`. \\
`ax` is the axes object and we can use it to set the labels and title of the plot.
`fig` is the figure object and we can use it to save the figure.\\
You can plot multiples axes in the same figure using `plt.subplots` and then plotting on the different axes. let's plot sin(x) and cos(x) in the same figure.


```python
import matplotlib.pyplot as plt
fig, ax = plt.subplots(nrows=2, ncols=1, figsize=(10,10))
x = np.linspace(0, 2*np.pi, 100) 
y1 = np.sin(x)
y2 = np.cos(x)
ax[0].plot(x, y1)
ax[1].plot(x, y2)
ax[0].set_xlabel('x')
ax[0].set_ylabel('sin(x)')
ax[1].set_xlabel('x')
ax[1].set_ylabel('cos(x)')
plt.tight_layout()
plt.show()
```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

**Assignment 3.1** Plot the function $y = cos(x)+cos(3x)$ from 0 to $6\pi$ and add below the plot, it's fourier transform. Use the variables `y` and `x` to define $y$ and $x$.

**TIP** : Use the `np.fft.fft` function to calculate the fourier transform and use the `np.fft.fftfreq` function to get the frequencies. The `fftfreq` gives you the x-axis of the fourier transform and the `fft` gives you the y-axis. Check the documentation of the functions to understand how they work and what the input parameters are.

Save the value of fftfreq in a variable `freq` and the value of fft in a variable `y_fft` and plot `y_fft` against `freq` in the same figure as the original plot. Notice that the values returned value of `np.fft.fft` are complex numbers, you can use `np.abs` to get the magnitude of the complex numbers. You will also need to use `np.fft.fftshift` to shift the zero frequency to the center of the plot. Apply `np.fft.fftshift` on `freq` and `y_fft` before plotting them.

**Note** The FFT is a sum over all N samples. Therefore, increasing the number of samples increases the magnitude of the FFT even though the signal itself has not changed. Dividing the FFT by N normalizes this dependence on the number of samples, so the Fourier amplitudes represent the signal amplitudes rather than the number of samples used. To compute N use `N = len(y)`.
    

<!-- grade_id: cell-4b3264f5be3fb7e7, points: 0, grade: false, solution: true -->
```python
fig,ax = None, None
freq = None
y_fft = None
# YOUR CODE HERE
raise NotImplementedError()
plt.tight_layout()
plt.show()
```

<!-- grade_id: cell-a6336c9e88e5eed1, points: 0, grade: true, solution: false -->
```python
assert len(x) == len(freq), "The frequency vector should have an equal length as x"
assert np.isclose(max(x), 6*np.pi, rtol=0.1), "Does x go to 6*pi?"
assert np.isclose(max(y), 2, rtol=0.1), "Did you compute y correctly?"
assert freq[int(len(x)/2)] == 0, "Did you use fftshift?"
```

<div style="border-left: 1px solid black; padding: 1em; margin: 1em 0;">

**Assignment 3.2:** In this assignment, you will plot Lissajous curves on a 3x3 grid, each with different parameters (order). Lissajous curves are graphs of a system of parametric equations describing harmonic motion. The simplified curves are given by the following equations:

$$
x(t) = \sin(a t + \delta)
$$
$$
y(t) = \sin(b t)
$$

where:


- $ a $ and $ b $ determine the number of lobes of the curve.
- $ \delta $ is the phase difference.
    
    For more information on Lissajous Figures: https://en.wikipedia.org/wiki/Lissajous_curve

### Problem Statement:

1. **Create a function** named `lissajous_curve` that generates the x and y coordinates of a Lissajous curve. The function should take the following parameters:
   - `a`: Frequency in the x direction.
   - `b`: Frequency in the y direction.
   - `t`: Time array (a NumPy array).
   - `delta`: phase difference
   
   The function should return the x and y coordinates as NumPy arrays.

2. **Create a 3x3 grid of plots** where each subplot displays a Lissajous curve with different parameters. Use the following orders (a, b) and a fixed phase difference ($\delta = \frac{\pi}{2}$):

   - Row 1: \( (a, b) \) = (1, 1), (1, 2), (1, 3)
   - Row 2: \( (a, b) \) = (2, 1), (2, 2), (2, 3)
   - Row 3: \( (a, b) \) = (3, 1), (3, 2), (3, 3)

3. **Customize your plots**:
   - Label the x-axis as "x(t)" and the y-axis as "y(t)" for each subplot.
   - Set the title of each subplot to reflect the order (e.g., "a=1, b=1").

4. **Display the entire figure** with the 3x3 grid of Lissajous curves.

**TIP:** Use NumPy to generate the time array and the sine functions. Use Matplotlib for plotting.


---

### Solution Outline:

1. **Define the `lissajous_curve` function**:
   ```python
   import numpy as np
   import matplotlib.pyplot as plt

   def lissajous_curve(a, b, delta, t):
       x = ** ** **
       y = np.sin(b * t)
       return x, y

   t = np.linspace(0, 2 * np.pi, 1000)

   fig, ax = plt.subplots(nrows=3, ncols=3,figsize=(20,20))
   for i in range(3):
      for j in range(3):
         a, b = i+1, j+1 
         x,y = lissajous_curve(a,b, np.pi/2)
         ax[i, j].plot(x, y)
   plt.show()

   ``` 


<!-- grade_id: cell-e0b55d988e8614bd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

There are thousands of settings to customize your figures. But that would take us too far. You can always use Google or ChatGPT to find better ways to plot your functions.

<div class="alert alert-block alert-info">
<center><b>This concludes the first introduction session, we hope you learned a thing or two. </b></center></div>